# Basic linear algebra in Julia
Author: Andreas Noack Jensen (MIT & JuliaComputing) (https://twitter.com/anoackjensen?lang=en)
(with edits from Jane Herriman)

First let's define a random matrix

```julia
A = rand(1:4,3,3)
```

Output:
```
3×3 Matrix{Int64}:
 2  4  3
 3  4  1
 3  1  2
```

Define a vector of ones

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages
#### Multiplication

```julia
b = A*x
```

Output:
```
3-element Vector{Float64}:
 9.0
 8.0
 6.0
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 2  3  3
 4  4  1
 3  1  2
```

and we can get the transpose with

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 2  3  3
 4  4  1
 3  1  2
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 22  23  15
 23  33  18
 15  18  14
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ function.

```julia
A\b
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.97379    0.0261919
 0.542349   0.30641
 0.0833862  0.576553
```

```julia
Atall\b
```

Output:
```
2-element Vector{Float64}:
 9.10705853153404
 9.280826926155607
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankdef = hcat(v, v)
```

Output:
```
3×2 Matrix{Float64}:
 0.521319  0.521319
 0.457765  0.457765
 0.327194  0.327194
```

```julia
rankdef\b
```

Output:
```
2-element Vector{Float64}:
 8.767453040927371
 8.767453040927371
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.194555  0.294871  0.17081
 0.682405  0.702613  0.445114
```

```julia
Ashort\bshort
```

Output:
```
3-element Vector{Float64}:
 -9.906603147004386
  8.602855646536772
  2.2693882184279075
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`.



```julia
v = [1,2,3]
dot_v = v'v
```

Output:
```
14
```

```julia
map(x -> x^2, v)
v1 = [1, 2, 3]
v2 = [2, 3, 4]
v1 .* v2

[i * j for i in v1, j in v2] # compare to a list apprehension

```

Output:
```
3×3 Matrix{Int64}:
 2  3   4
 4  6   8
 6  9  12
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia
outer_v = [i * j for i in v, j in v] # using list apprehession 
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 2  4  6
 3  6  9
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

#### 10.3 
Use [LinearAlgebra.cross](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/#LinearAlgebra.cross) to compute the cross product of a vector v with itself and assign it to variable `cross_v`

```julia
using LinearAlgebra
cross_v = cross(v, v)
```

Output:
```
3-element Vector{Int64}:
 0
 0
 0
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```