# Loops

Topics:
1. `while` loops
2. `for` loops
<br>

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```

For example, we could use `while` to count or to iterate over an array.

```julia
n = 0
while n < 10
    n += 1
    println(n)
end
n
```

Output:
```
1
2
3
4
5
6
7
8
9
10
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

i = 1
while i <= length(myfriends)
    friend = myfriends[i]
    println("Hi $friend, it's great to see you!")
    i += 1
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

We could use a for loop to generate the same results as either of the examples above:

```julia
for n in 1:10
    println(n)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```julia
# note the syntax of in can change
for n = 1:5
    println(n)
end
for n ∈ 5:10
    println(n)
end
```

Output:
```
1
2
3
4
5
5
6
7
8
9
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

for friend in myfriends
    println("Hi $friend, it's great to see you!")
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

Now let's use `for` loops to create some addition tables, where the value of every entry is the sum of its row and column indices. <br>

Note that we iterate over this array via column-major loops in order to get the best performance. More information about fast indexing of multidimensional arrays inside nested loops can be found at https://docs.julialang.org/en/v1/manual/performance-tips/#Access-arrays-in-memory-order,-along-columns-1

First, we initialize an array with zeros.

```julia
m, n = 5, 5
A = fill(0, (m, n))
```

Output:
```
5×5 Matrix{Int64}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
fill("ha", (4, 5))
```

Output:
```
4×5 Matrix{String}:
 "ha"  "ha"  "ha"  "ha"  "ha"
 "ha"  "ha"  "ha"  "ha"  "ha"
 "ha"  "ha"  "ha"  "ha"  "ha"
 "ha"  "ha"  "ha"  "ha"  "ha"
```

```julia
B = fill(0, (5, 5))
```

Output:
```
5×5 Matrix{Int64}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
x = 0
for j in 1:n
    for i in 1:m
        x += 1
        A[i, j] = i + j
        B[i, j] = x
    end
end
A
B
```

Output:
```
5×5 Matrix{Int64}:
 1   6  11  16  21
 2   7  12  17  22
 3   8  13  18  23
 4   9  14  19  24
 5  10  15  20  25
```

Here's some syntactic sugar for the same nested `for` loop

```julia
B = fill(0, (m, n))
```

Output:
```
5×5 Matrix{Int64}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
# row seems to be looking through first so the second argument
x = 0
for j in 1:n, i in 1:m
    x += 1
    B[i, j] = x
end
B
```

Output:
```
5×5 Matrix{Int64}:
 1   6  11  16  21
 2   7  12  17  22
 3   8  13  18  23
 4   9  14  19  24
 5  10  15  20  25
```

```julia
x = 0
for i in 1:m, j in 1:n
    x += 1
    B[i, j] = x
end
B # after, is your child, your nested iteractor that will loop through rist
```

Output:
```
5×5 Matrix{Int64}:
  1   2   3   4   5
  6   7   8   9  10
 11  12  13  14  15
 16  17  18  19  20
 21  22  23  24  25
```

The more "Julia" way to create this addition table would have been with an *array comprehension*.

```julia
C = [i + j for i in 1:m, j in 1:n]
```

Output:
```
5×5 Matrix{Int64}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

### Exercises

#### 4.1 
Loop over integers between 1 and 100 and print their squares.

```julia
for i in 1:100
    println(i)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
20
21
22
23
24
25
26
27
28
29
30
31
32
33
34
35
36
37
38
39
40
41
42
43
44
45
46
47
48
49
50
51
52
53
54
55
56
57
58
59
60
61
62
63
64
65
66
67
68
69
70
71
72
73
74
75
76
77
78
79
80
81
82
83
84
85
86
87
88
89
90
91
92
93
94
95
96
97
98
99
100
```

#### 4.2 
Add to the code above a bit to create a dictionary, `squares` that holds integers and their squares as key, value pairs such that

```julia
squares[10] == 100
```

```julia
squares = Dict(i => i^2 for i in 10:11)
```

Output:
```
Dict{Int64, Int64} with 2 entries:
  11 => 121
  10 => 100
```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia
squares_arr = [(10*(j -1) + i)^2 for i = 1:10, j in 1:10]
```

Output:
```
10×10 Matrix{Int64}:
   1  121  441   961  1681  2601  3721  5041  6561   8281
   4  144  484  1024  1764  2704  3844  5184  6724   8464
   9  169  529  1089  1849  2809  3969  5329  6889   8649
  16  196  576  1156  1936  2916  4096  5476  7056   8836
  25  225  625  1225  2025  3025  4225  5625  7225   9025
  36  256  676  1296  2116  3136  4356  5776  7396   9216
  49  289  729  1369  2209  3249  4489  5929  7569   9409
  64  324  784  1444  2304  3364  4624  6084  7744   9604
  81  361  841  1521  2401  3481  4761  6241  7921   9801
 100  400  900  1600  2500  3600  4900  6400  8100  10000
```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```